Activation of the P2RX7/IL-18 pathway in immune cells attenuates lung fibrosis

Author Response

The following is the authors’ response to the previous reviews.

Point to point response for the editors

We are deeply grateful for the time you have devoted to reviewing this manuscript, and we sincerely thank you. Your insightful feedback has been instrumental in enhancing the quality of our work.

In the revised version of the manuscript, we have carefully addressed each of the concerns you raised. Below, you will find a detailed summary of how your feedback has been incorporated to improve the overall content and clarity of the document.

1. P2RX7 effects: In Figure 2, the vehicle treated P2RX7 knockout (panel M) shows an Ashcroft score of about 1.5 after BLM. Comparing this to the Ashcroft score of 3 after BLM in the wildtype (panel C) suggests that P2RX7 deletion is an effective way to reduce fibrosis by half!.

The argument that HEI3090 also reduces fibrosis by activating P2RX7 is of course very difficult to convey and it seems contradictory that P2RX7 deletion and P2RX7 activation can be both anti-fibrotic. This is an unusual claim and confuses the reviewers as well as the future readers.

This has many important health implications because activating an inflammatory pathway via P2RX7 and IL-18 could be risky in terms of a fibrosis treatment as inflammatory activation can also worsen fibrosis. The authors' own P2RX7 KO data (untreated vehicle groups) indeed confirms that P2RX7 can be pro-fibrotic.

We thank the editors for their comment highlighting the lack of clarity in our message. Indeed, we verified whether the antifibrotic action of HEI3090 depends on the expression of P2RX7 by inducing lung fibrosis in P2RX7 KO mice. In doing so, we initially observed that P2RX7 plays a role in the development of BLM-induced lung fibrosis. This is illustrated by a decrease of 50% in the Ashcroft score, as shown in Figure 2M and Supplemental Figure 2C of the revised manuscript.

To increase the clarity of your message, we added in the text the following paragraph:

"We further verified whether the antifibrotic action of HEI3090 depends on the expression of P2RX7 by inducing lung fibrosis in p2rx7 knockout (KO) mice. In doing so, we initially observed that P2RX7 plays a role in the development of BLM-induced lung fibrosis. This is illustrated by a decrease of 50% in the Ashcroft score, with a mean value of 1.7 in P2RX7 knockout mice compared to 3 in wild-type mice (Figure 2M and Supplemental Figure 2C). It is important to note that p2rx7 -/- mice still exhibit signs of lung fibrosis, such as thickening of the alveolar wall and a reduction in free air space, in comparison to naïve mice that received PBS instead of BLM (see Supplemental Figure 2A). This result confirms a previous report indicating that BLM-induced lung fibrosis partially depends on the activation of the P2RX7/pannexin-1 axis, leading to the production of IL-1β in the lung. Additionally, in contrast to the observations in WT mice, HEI3090 failed to attenuate the remaining lung fibrosis in p2rx7 -/- mice, as measured by the Ashcroft score (Figure 2M), the percentage of lung tissue with fibrotic lesions, or the intensity of collagen fibers (Supplemental Figure 2D). These results show that P2RX7 alone participates in fibrosis and that HEI3090 exerts a specific antifibrotic effect through this receptor (see Supplemental Figure 2C)."

Since we used the HEI3090 compound in this study and to be closer to the results, we have replaced the title of 2 chapters in the results section as followed:

“HEI3090 inhibits the onset of pulmonary fibrosis in the bleomycin mouse model” instead of P2RX7 activation inhibits the onset of pulmonary fibrosis in the bleomycin mouse model and “HEI3090 shapes immune cell infiltration in the lungs" instead of P2RX7 activation shapes immune cell infiltration in the lungs

We concur that the observation of both anti-fibrotic effects following P2RX7 deletion and P2RX7 activation appears contradictory. This specific aspect has been thoroughly addressed and extensively discussed in the revised manuscript.

“A major unmet need in the field of IPF is new treatment to fight this uncurable disease. In this preclinical study, we demonstrate the ability of immune cells to limit lung fibrosis progression. Based on the hypothesis that a local activation of a T cell immune response and upregulation of IFN-γ production has antifibrotic proprieties, we used the HEI3090 positive modulator of the purinergic receptor P2RX7, previously developed in our laboratory (Douguet et al., 2021), to demonstrate that activation of the P2RX7/IL-18 pathway attenuates lung fibrosis in the bleomycin mouse model. We have demonstrated that lung fibrosis progression is inhibited by HEI3090 in the fibrotic phase but also in the acute phase of the BLM fibrosis mouse model, i.e. during the period of inflammation. This lung fibrosis mouse model commonly employed in preclinical investigations, has recently been recognized as the optimal model for studying IPF (Jenkins et al., 2017). In this model, the intrapulmonary administration of BLM induces DNA damage in alveolar epithelial type 1 cells, triggering cellular demise and the release of ATP. The extracellular release of ATP from injured cells activates the P2RX7/pannexin 1 axis, initiating the maturation of IL1β and subsequent induction of inflammation and fibrosis. In line with this, mice lacking P2RX7 exhibited reduced neutrophil counts in their bronchoalveolar fluids and decreased levels of IL1β in their lungs compared to WT mice (Riteau et al., 2010). Based on these findings, Riteau and colleagues postulated that the inhibition of P2RX7 activity may offer a potential strategy for the therapeutic control of fibrosis in lung injury. In the present study we provided strong evidence showing that selective activation of P2RX7 on immune cells, through the use of HEI3090, can dampen inflammation and fibrosis by releasing IL-18. The efficacy of HEI3090 to inhibit lung fibrosis was evaluated histologically on the whole lung’s surface by evaluating the severity of fibrosis using three independent approaches applied to the whole lung, the Ashcroft score, quantification of fibroblasts/myofibroblasts (CD140a) and polarized-light microscopy of Sirius Red staining to quantify collagen fibers. All these methods of fibrosis assessment revealed that HEI3090 exerts an inhibitory effect on lung fibrosis, underscoring the necessity for a thorough pre-clinical assessment of HEI3090's mode of action. Notably, HEI3090 functions as an activator, rather than an inhibitor, of P2RX7, further emphasizing the importance of elucidating its intricate mechanisms.”

We trust that the detailed explanation provided therein will adequately persuade both the reviewers and future readers.

1. The statistical concerns are based on the phrasing of "the experiment was stopped when significantly statistical results were observed". This is different from the power analysis approach that the authors describe in their latest rebuttal. However, it raises the question why the power analysis was performed using "on a one-way ANOVA analysis comparing in each experiment the vehicle and the treated group". The analyses in the manuscript use the Mann-Whitney test for several comparisons which ahs the assumption that the samples do NOT have a normal distribution. An ANOVA and t-tests have the assumption that samples are normally distributed. If the power analysis and "statistical forecasting" assumed a normal distribution and used an ANOVA, then shouldn't all the analyses also use a statistical test appropriate for normally distributed samples such as ANOVA and t-tests?

Several of the data points in the figures seem to be normally distributed and therefore t-test for two group comparisons would be more appropriate. The most rigorous approach would be to check for normal distribution before choosing the correct statistical test and using the t-test/ANOVA in normally distributed data as well as Mann-Whitney for non-normally distributed data.

We described in the Material and Method section of the revised manuscript our approach to determine the size of experimental group.

“The determination of experimental group sizes involved conducting a pilot experiment with four mice in each group. Subsequently, a power analysis, based on the pilot experiment's findings (which revealed a 40% difference with a standard error of 0.9, α risk of 0.05, and power of 0.8), was performed to ascertain the appropriate group size for studying the effects of HEI3090 on BLM-induced lung fibrosis. The results of the pilot experiment and power analysis indicated that a group size of four mice was sufficient to characterize the observed effects. For each full-scale experiment, we initiated the study with 6 to 8 mice per group, ensuring a minimum of 5 mice in each group for robust statistical analysis. Additionally, we systematically employed the ROULT method to identify and subsequently exclude any outliers present in each experiment before conducting statistical analyses”.

We now described in the Material and Method section how we carried out the statistical analyses.

“Quantitative data were described and presented graphically as medians and interquartiles or means and standard deviations. The distribution normality was tested with the Shapiro's test and homoscedasticity with a Bartlett's test. For two categories, statistical comparisons were performed using the Student's t-test or the Mann–Whitney's test. For three and more categories, analysis of variance (ANOVA) or non-parametric data with Kruskal–Wallis was performed to test variables expressed as categories versus continuous variables. If this test was significant, we used the Tukey's test to compare these categories and the Bonferroni’s test to adjust the significant threshold. For the Gene Set Enrichment Analyses (GSEA), bilateral Kolmogorov–Smirnov test, and false discovery rate (FDR) were used. All statistical analyses were performed by biostatistician using Prism8 program from GraphPad software. Tests of significance was two-tailed and considered significant with an alpha level of P < 0.05. (graphically: * for P < 0.05, ** for P < 0.01, *** for P < 0.001).”

We also added in the legend of each figure, the statistical analysis used to determine each p-values.

1. Adoptive transfer: The concerns of the reviewers include an unclear analysis of the effects of adoptive transfer itself and the approaches used to analyze the data independent of the HEI3090 effect. For example, in Figure 4, the adoptive transfer IL18-/- cells (vehicle group) leads to an Ashcroft score of about 1 and among the lowest of the BLM exposed mice. Does that mean that IL18 is pro-fibrotic and that its absence is beneficial? If yes, it would go against the core premise of the study that IL18 is beneficial. Statistical comparisons of the all the vehicle conditions in the adoptive transfer would help clarify whether adoptive transfer of NLRP3-/-, IL18-/- in wild-type and P2RX7-/- mice reduces or increases fibrosis. Such multiple comparisons are necessary to fully understand the adoptive transfer studies and would also require the appropriate statistical test with corrections for multiple comparisons such as Kruskal-Wallis for data without normal distribution and ANOVA with post hoc correction for normal distribution.

We added a new paragraph in the revised version of the manuscript to explain the adoptive transfer approach.

“We wanted to further investigate the mechanism of action of HEI3090 by identifying the cellular compartment and signaling pathway required for its activity. Since the expression of P2RX7 and the P2RX7-dependent release of IL-18 are mostly associated with immune cells (Ferrari et al., 2006), and since HEI3090 shapes the lung immune landscape (Figure 3), we investigated whether immune cells were required for the antifibrotic effect of HEI3090. To do so, we conducted adoptive transfer experiments wherein immune cells from a donor mouse were intravenously injected one day before BLM administration into an acceptor mouse. The intravenous injection route was chosen as it is a standard method for targeting the lungs, as previously documented (Wei and Zhao, 2014). This approach was previously used with success in our laboratory (Douguet et al., 2021). It is noteworthy that this adoptive transfer approach did not influence the response to HEI3090. This was observed consistently in both p2rx7 -/- mice and p2rx7 -/- mice that received splenocytes of the same genetic background. In both cases, HEI3090 failed to mitigate lung fibrosis, as depicted in Figure 2M and Supplemental Figures 2D and 6A and B.”

We added the Supplemental Figure 7 showing that the genetic background does not impact lung fibrosis at steady step levels where p-values were analyzed by one-way ANOVA, with Kruskal-Wallis test for multiple comparisons.

Supplemental Figure 7 : The genetic background does not impact lung fibrosis at steady step levels.

p2rx7-/- mice were given 3.106 WT, nlrp3-/ , i118-/ or illb -l- splenocytes i_v_ one day prior to BLM delivery (i_n_ 2.5 LJ/kg)_ p2rx7-/- mice or p2rx7-/- mice adoptively transferred with splenocytes from indicated genetic background were treated daily i.p_ with mg/kg HE13090 or vehicle for 14 days. Fibrosis score assessed by the Ashcroft method. P-values were analyzed on all treated and non treated groups by one-way ANOVA, with Kruskal-Wallis test for multiple comparisons. The violin plot illustrates the distribution of Ashcroft scores across indicated experimental groups. The width of the violin at each point represents the density of data, and the central line indicates the median expression level. Each point represents one biological replicate. ns, not significant

Changes from the earlier version to the revised version

Abstract

Abstract

Idiopathic pulmonary fibrosis (IPF) is an aggressive interstitial lung disease associated with progressive and irreversible deterioration of respiratory functions that lacks curative therapies. Despite IPF being associated with a dysregulated immune response, current antifibrotics aim only at limiting fibroproliferation. We show here that the P2RX7/IL-18/IFNG axis is downregulated in IPF patients and that P2RX7 has immunoregulatory functions. Using our positive modulator of P2RX7, we show that activation of the P2RX7/IL-18 axis in immune cells limits lung fibrosis progression in a mouse model by favoring an anti-fibrotic immune environment, with notably an enhanced IL-18-dependent IFN-γ production by lung T cells leading to a decreased production of IL-17 and TGFβ. Overall, we show the ability of the immune system to limit lung fibrosis progression by targeting the immunomodulator P2RX7. Hence, treatment with a small activator of P2RX7 may represent a promising strategy to help patients with lung fibrosis.

Introduction

Idiopathic pulmonary fibrosis (IPF) is an aggressive interstitial lung disease associated with progressive deterioration of respiratory function that is ultimately fatal. It is characterized by destruction of the lung architecture due to accumulation of fibroblasts and extracellular matrix proteins, resulting in increased lung stiffness and impaired normal breathing.

Earlier version: Pirfenidone and nintedanib have been FDA approved for the treatment of IPF since 2014. They target respectively the key fibrotic cytokine TGFβ and several receptor tyrosine kinases, thereby affecting fibroblast activation and extracellular matrix protein production [1]. However, they only slow down the progression of the disease, so new therapeutic strategies and targets are needed.

Revised version: Pirfenidone and nintedanib have been FDA approved for the treatment of IPF since 2014. They target respectively the key fibrotic cytokine TGFβ and several receptor tyrosine kinases, thereby affecting fibroblast activation and extracellular matrix protein production (Heukels et al., 2019). However, they only slow down the progression of the disease, so new therapeutic strategies and targets are needed.

Earlier version: Fibrosis is also associated with inflammation. In fact, fibrosis is a process of excessive wound healing and tissue remodeling due to repeated epithelial injuries releasing damage-associated molecular patterns (DAMPs) that trigger both the adaptive and innate immune systems. Although inflammation has not been considered a target in IPF, due to unsuccessful initial clinical trials of anti-inflammatory drugs [2] or cyclophosphamide during exacerbations [3], growing evidence suggest that altering specific immune populations that promote or attenuate disease progression may be beneficial [4].

Revised version: Fibrosis is also associated with inflammation. In fact, fibrosis is a process of excessive wound healing and tissue remodeling due to repeated epithelial injuries releasing damage-associated molecular patterns (DAMPs) that trigger both the adaptive and innate immune systems. Although inflammation has not been considered a target in IPF, due to unsuccessful initial clinical trials of anti-inflammatory drugs (Idiopathic Pulmonary Fibrosis Clinical Research Network et al., 2012) or cyclophosphamide during exacerbations (Naccache et al., 2022), growing evidence suggest that altering specific immune populations that promote or attenuate disease progression may be beneficial (Shenderov et al., 2021).

Earlier version: Extracellular adenosine triphosphate (eATP) is a DAMP, released in high concentrations from injured cells in IPF patients [5]. High levels of eATP are recognized by the P2X7 receptor (P2RX7) and are both required for the establishment of the bleomycin lung fibrosis mouse model [5]. Activation of P2RX7 by ATP [6] leads to conformational rearrangements allowing influx of Ca2+ and Na+ ions into and the efflux of K+ ions from the cell [7] and to the formation of a large pore mediating macrophage cell death. The nature of this large pore was subjected to controverse until the recent discovery that the macropore function is intrinsic to P2RX7 and does not involve progressive pore dilatation [8,9]. Another feature of P2RX7 is its ability to activate the NLRP3 inflammasome, an event linked to potassium efflux [10], which requires both P2RX7 and TWIK2 activation [11] and leads to the release of mature IL-β and IL-18 [12], through gasdermin D pores. Consequently, P2RX7 has the ability to trigger an immune response [13].

Revised version: Extracellular adenosine triphosphate (eATP) is a DAMP, released in high concentrations from injured cells in IPF patients (Riteau et al., 2010). High levels of eATP are recognized by the P2X7 receptor (P2RX7) and P2RX7 participates in the establishment of the bleomycin lung fibrosis mouse model (Riteau et al., 2010). Extracellular activation of P2RX7 by ATP (Surprenant et al., 1996) leads to conformational rearrangements allowing influx of Ca2+ and Na+ ions into and the efflux of K+ ions from the cell (Karasawa et al., 2017) and to the formation of a large pore mediating macrophage cell death. The nature of this large pore was subjected to controverse until the recent discovery that the macropore function is intrinsic to P2RX7 and does not involve progressive pore dilatation (Harkat et al., 2017; Pippel et al., 2017). Another feature of P2RX7 is its ability to activate the NLRP3 inflammasome, an event linked to potassium efflux (Di et al., 2018; Muñoz-Planillo et al., 2013), which requires both P2RX7 and TWIK2 activation(Di et al., 2018) and leads to the release of mature IL-1β and IL-18 (Perregaux et al., 2000), through gasdermin D pores. Consequently, P2RX7 has the ability to trigger an immune response (Janho Dit Hreich et al., 2023).

Earlier version: IL-β is a proinflammatory cytokine with high profibrotic properties, as it promotes collagen deposition through IL-17A and TGFβ production[14–16] but also promotes the activation and recruitment of inflammatory cells such as eosinophils and neutrophils. Indeed, deficiency of IL-1βR or its blockade ameliorate experimental fibrosis [17–19]. In contrast, the role of IL-18 is not clear. Indeed, conflicting experimental studies show that IL-18 could either promote [20] or attenuate [21] fibrosis. However, high levels of IL-18BP, a natural antagonist of IL-18, are associated with reduced overall survival in IPF patients [22], suggesting that the activity of IL-18 may be required for improved survival or alternatively that high levels of IL-18 are accompanied by more IL-18BP, which leads to poorer outcomes.

Revised version: IL-1β is a proinflammatory cytokine with high profibrotic properties, as it promotes collagen deposition through IL-17A and TGFβ production (Doerner and Zuraw, 2009; Sutton et al., 2009; Wilson et al., 2010) but also promotes the activation and recruitment of inflammatory cells such as eosinophils and neutrophils. Indeed, deficiency of IL-1βR or its blockade ameliorate experimental fibrosis (Couillin et al., 2009; Gasse et al., 2009, 2007). In contrast, the role of IL-18 is not clear. Indeed, conflicting experimental studies show that IL-18 could either promote (Zhang et al., 2019) or attenuate (Nakatani-Okuda et al., 2005) fibrosis. However, high levels of IL-18BP, a natural antagonist of IL-18, are associated with reduced overall survival in IPF patients (Nakanishi et al., 2021), suggesting that the activity of IL-18 may be required for improved survival or alternatively that high levels of IL-18 are accompanied by more IL-18BP, which leads to poorer outcomes.

Earlier version: IL-18 was originally described as IFN-γ-inducing factor (IGIF) and therefore boosts IFN-γ production by T cells and NK cells [23]. Not only has IFN-γ antiproliferative properties [24] but it also inhibits TGFβ activity [25,26] and therefore inhibits fibroblast activation and differentiation into myofibroblasts, alleviates TGFβ-mediated immunosuppression, inhibits extracellular matrix accumulation and collagen production [27–30] and thus promotes an antifibrotic immune microenvironment, making IFN-γ a cytokine with antifibrotic properties. However, parenteral systemic administration of IFN-γ failed in clinical trials (NSPIRE; NCT 00075998) [31], whereas local administration by inhalation showed promising results [32–36].

Revised version: IL-18 was originally described as IFN-γ-inducing factor (IGIF) and therefore boosts IFN-γ production by T cells and NK cells (Okamura et al., 1995). Not only has IFN-γ antiproliferative properties (Elias et al., 1987) but it also inhibits TGFβ activity (Ulloa et al., 1999; Varga et al., 1990) and therefore inhibits fibroblast activation and differentiation into myofibroblasts, alleviates TGFβ-mediated immunosuppression, inhibits extracellular matrix accumulation and collagen production(Duncan and Berman, 1985; Gillery et al., 1992; Gurujeyalakshmi and Giri, 1995; Yuan et al., 1999) and thus promotes an antifibrotic immune microenvironment, making IFN-γ a cytokine with antifibrotic properties. However, parenteral systemic administration of IFN-γ failed in clinical trials (INSPIRE; NCT 00075998) (King et al., 2009), whereas local administration by inhalation showed promising results (Diaz et al., 2012; Fusiak et al., 2015; Prior and Haslam, 1992; Skaria et al., 2015; Smaldone, 2018).

Earlier version: One way to increase IFN-γ production locally and selectively in the lung would be to alter the phenotype of T cells since the polarization of T lymphocytes has been shown to impact fibroblasts’ fate and immune infiltrate [4,37]. Indeed, T cells have been recently shown to selectively kill myofibroblasts through IFN-γ release and limit the progression of lung and liver fibrosis in preclinical models [38] and set up an immune memory in the long term since IFN-γ producing tissue resident memory T cells protect against fibrosis progression [39], highlighting the importance IFN-γ producing T cells in this disease. Accordingly, CD4‘-producing IFN-γ T cells are decreased in IPF and correlate with a better prognosis in IPF patients [37,40,41].

Revised version: One way to increase IFN-γ production locally and selectively in the lung would be to alter the phenotype of T cells since the polarization of T lymphocytes has been shown to impact fibroblasts’ fate and immune infiltrate (Luzina et al., 2008; Shenderov et al., 2021). Indeed, T cells have been recently shown to selectively kill myofibroblasts through IFN-γ release and limit the progression of lung and liver fibrosis in preclinical models (Sobecki et al., 2022) and set up an immune memory in the long term since IFN-γ-producing tissue resident memory T cells protect against fibrosis progression (Collins et al., 2016), highlighting the importance IFN-γ producing T cells in this disease. Accordingly, CD4+-producing IFN-γ T cells are decreased in IPF and correlate with a better prognosis in IPF patients (Giri et al., 1986; Luzina et al., 2008; Xu et al., 2006).

Added in the revised version: In light of IL-18’s demonstrated capability to modulate the T cell phenotype by eliciting IFN-γ, we have posited a novel therapeutic approach for pulmonary fibrosis. Our proposal involves augmenting local IFN-γ production through targeted intervention in the P2RX7/IL-18 axis. This strategy aims to harness the immunomodulatory potential of this axis, offering a promising avenue for therapeutic intervention in pulmonary fibrosis. To implement this strategy, we utilized HEI3090, a chemical compound developed in our laboratory. HEI3090 has been previously reported to enhance the channel activity and large pore opening of P2RX7 and also to increase IL-18 levels (but not IL-1β) upon activation of the P2RX7/NLRP3 pathway (Douguet et al., 2021).

Removed in the revised version: Given the ability of IL-18 to shape the phenotype of T cells by inducing IFN-γ, we proposed to increase local IFN-γ production via the P2RX7/IL-18 axis as a therapeutic strategy in pulmonary fibrosis. We used a P2RX7-specific positive modulator, developed in our laboratory, which has the particularity to increase IL-18 levels in the presence of high eATP [42].

Results

Earlier version: Expression of P2RX7 and IL-18 activity are dampened in IPF patients

Revised version: Expression of P2RX7 and IL-18 activity are downregulated in IPF patients

Earlier version: The canonical release of IL-18 is due to activation of the P2RX7/NLRP3 pathway [12]. Since high levels of eATP are found in IPF patients [5] and P2RX7 is activated by such levels, it was of particular interest to investigate the involvement of P2RX7 in this disease. We used a publicly available dataset of lung homogenates from control and IPF patients and compared the mRNA expression levels of P2RX7 and markers of fibrosis, namely ACTA2, COLIA2, COL3A1 and TGFB3. We found that the mRNA expression of P2RX7 is downregulated in IPF patients (Figure 1A and B), as well as the components of the NLRP3 inflammasome (Supplemental Figure 1). Since IL-18 is constitutively expressed [43], which partly explains the lack of difference between control and IPF patients (Figure 1B), we investigated the signaling pathway downstream of IL-18. IL-18 binds to its receptor IL-18R1 coupled to its adaptor protein IL-18RAP which is required for IL-18 signaling and IFN-γ expression. We showed that IL-18R1, IL-18RAP and IFN-γ (Figure 1B) are downregulated in IPF patients. Knowing that the modulation of the phenotype of T cells is promising in IPF [4], we checked whether P2RX7 and IL-18 are linked to an immune response in IPF patients using Gene Set Enrichment Analyses. Indeed, we showed that the expression of P2RX7 and IL-18 signaling (IL-18 and IL- 18RAP) correlates with the IFN-γ response as well as with immunoregulatory interactions required for changing the phenotype of T cells (Figure 1C, Supplemental Figure 1B). This transcriptomic analysis highlights that the P2RX7/IL-18 signaling pathway is dampened and suggests that this pathway may be able to modulate the immune response in IPF patients.

Revised version: The canonical release of IL-18 is due to activation of the P2RX7/NLRP3 pathway (Perregaux et al., 2000). Since high levels of eATP are found in IPF patients (Riteau et al., 2010) and P2RX7 is activated by such levels, it was of particular interest to investigate the involvement of P2RX7 in this disease. We used a publicly available dataset of lung homogenates from control and IPF patients and compared the mRNA expression levels of P2RX7 and markers of fibrosis, namely ACTA2, COL1A2, COL3A1 and TGFB3. We found that the mRNA expression of P2RX7 is downregulated in IPF patients (Figure 1A and B), as well as the components of the NLRP3 inflammasome (Supplemental Figure 1). Since IL-18 is constitutively expressed (Puren et al., 1999), which partly explains the lack of difference between control and IPF patients (Figure 1B), we investigated the signaling pathway downstream of IL-18. IL-18 binds to its receptor IL-18R1 coupled to its adaptor protein IL-18RAP which is required for IL-18 signaling and IFN-γ expression. We showed that IL-18R1, IL-18RAP and IFN-γ (Figure 1B) are downregulated in IPF patients. Knowing that the modulation of the phenotype of T cells is promising in IPF (Shenderov et al., 2021), we checked whether P2RX7 and IL-18 are linked to an immune response in IPF patients using Gene Set Enrichment Analyses. Indeed, we showed that the expression of P2RX7 and IL-18 signaling (IL-18 and IL-18RAP) correlates with the IFN-γ response as well as with immunoregulatory interactions required for changing the phenotype of T cells (Figure 1C, Supplemental Figure 1B). This transcriptomic analysis highlights that the P2RX7/IL-18 signaling pathway is downregulated and suggests that this pathway may be able to modulate the immune response in IPF patients.

The P2RX7/IL-18/IFN-γ pathway is downregulated in IPF

Earlier version: (A) Heatmap of mRNA expression of P2RX7 in control and IPF patients with a cluster of fibrosis-associated genes. Raw p-values are shown (Limma). (B) mRNA expression of P2RX7, IL18, IL18R1, IL18RAP and IFNG between control and IPF patients from 213 individuals, corresponding to 91 controls and 122 IPF patients. Two-tailed unpaired t-test with Welch’s correction, ***p<0.001, ****P<0.0001. (C) Gene set enrichment analysis (GSEA) plot associating P2RX7 mRNA levels from IPF patients with three immunological signatures. The green line represents the enrichment score and the black lines the specific signature-associated genes. NES: Normalized entichment score, FDR: False discovery rate. Pearson's correlation test.

Revised version: (A) Heatmap of mRNA expression of P2RX7 in control and IPF patients with a cluster of fibrosis-associated genes. Raw p-values are shown (Limma). P-values were determined by Spearman test. (B) mRNA expression of P2RX7, IL18, IL18R1, IL18RAP and IFNG between control and IPF patients from 213 individuals, corresponding to 91 controls and 122 IPF patients. P-values were determined by Two-tailed unpaired t-test or Mann-Whitney t-test. ***p < 0.001, ****p < 0.0001. (C) Gene set enrichment analysis (GSEA) plot associating P2RX7 mRNA levels from IPF patients with three immunological signatures. The green line represents the enrichment score and the black lines the specific signature-associated genes. NES: Normalized enrichment score, p-values (bilateral Kolmogorov–Smirnov), and false discovery rate (FDR) are shown.

Earlier version: Activation of P2RX7 inhibits the onset of pulmonary fibrosis in the bleomycin mouse model

Revised version: HEI3090 inhibits the onset of pulmonary fibrosis in the bleomycin mouse model

Earlier version: Even though transcriptomic changes do not always reflect changes in the proteome these encouraging results led us to hypothesize that activation of the P2RX7/IL-18 signaling may restrain lung fibrosis progression. To test this hypothesis, we decided to boost the P2RX7/IL- 18 signaling in the bleomycin (BLM)-induced lung fibrosis mouse model. To do so, we used a positive modulator of P2RX7, called HEI3090, which enhances P2RX7’s activity only in the presence of high eATP levels [42]. Indeed, knowing that high eATP levels are found in IPF patients as well as in this mouse model of pulmonary fibrosis [5], we thought that HEI3090 will selectively enhance the activity of P2RX7 in the lungs.

Revised version: Even though transcriptomic changes do not always reflect changes in the proteome these encouraging results led us to hypothesize that activation of the P2RX7/IL-18 axis may restrain lung fibrosis progression by promoting an antifibrotic immune microenvironment. To test this hypothesis, we enhanced the P2RX7/IL-18 axis in the bleomycin (BLM)-induced lung fibrosis mouse model by using HEI3090, a chemical compound described to stimulate immune cells expressing P2RX7 to generate IL-18 in the presence of eATP (Douguet et al., 2021).

Earlier version: We first tested the antifibrotic potential of HEI3090 on mice having an established fibrosis (Figure 2A). Activation of P2RX7 with HEI3090 in mice 7 days after BLM administration reduced the development of pulmonary fibrosis, as evidenced by less thickening of alveolar walls and free air space (Figure 2B). Fibrosis severity was evaluated using the Ashcroft score. To overcome the heterogeneity of fibrosis within the lobes, we scored the whole surface of the lung, and the result represents the mean of each field (Figure 2C and supplementary Figure 2). This quantification showed that fibrosis is reduced by 40% in lungs of HEI-treated mice.

Revised version: We first tested the antifibrotic potential of HEI3090 on mice having an established fibrosis (Figure 2A). Activation of P2RX7 with HEI3090 in mice 7 days after BLM administration reduced the development of pulmonary fibrosis, as evidenced by less thickening of alveolar walls and free air space (Figure 2B). Fibrosis severity was evaluated using the Ashcroft score. To overcome the heterogeneity of fibrosis within the lobes, we scored the whole surface of the lung, and the result represents the mean of each field (Figure 2C and supplementary Figure 2 A and B). This quantification showed that fibrosis is reduced by 40% in lungs of HEI-treated mice.

Earlier version: Activation of P2RX7 with HEI3090 inhibits lung fibrosis progression

Revised version: HEI3090 inhibits lung fibrosis progression

Earlier version: (A) Experimental design. WT mice were given 2.5 U/kg of bleomycin by i.n. route. At the end of the inflammatory phase (d7), 1.5 mg/kg of HEI3090 or vehicle were given daily until day 21. (B) Representative images of lung sections at day 21 after treatment stained with H&E and Sirius Red, Scale bar= 100 µm. (C) Fibrosis score assessed by the Ashcroft method. (D) Quantification of fibroblasts/myofibroblasts in non-immune cell subset. (E) Representative images of lung sections at day 21 after treatment stained with Sirius Red, Scale bar= 100 µm. (F) Collagen levels in whole lung of mice assessed on Sirius Red-polarized images. (G) Experimental design. WT mice were given 2.5 U/kg of bleomycin by i.n. route. 1.5 mg/kg of HEI3090 or vehicle were given daily until day 14. (H) Representative images of lung sections at day 14 after treatment stained with H&E, Scale bar= 100 µm. (I) Fibrosis score assessed by the Ashcroft method. (J) Quantification of fibroblasts/myofibroblasts in non-immune cell subset. (K) Representative images of lung sections at day 14 after treatment stained with Sirius Red, Scale bar= 100 µm. (L) Collagen levels in whole lung of mice assessed on Sirius Red-polarized images. (M) P2rx7 KO mice were given 2.5 U/kg of bleomycin by i.n. route. 1.5 mg/kg of HEI3090 or vehicle were given daily until day 14. Fibrosis score assessed by the Ashcroft method is showed. Each point represents one mouse, two-tailed Mann-Whitney test, p values: *p<0.05, **p<0.01. WT: Wildtype, BLM: bleomycin, i.p.: intraperitoneal, 1.n.: intranasal, H&E: hematoxylin & eosin, AU: arbitrary units.

Revised version: (A) Experimental design. WT mice were given 2.5 U/kg of bleomycin by i.n. route. At the end of the inflammatory phase (d7), 1.5 mg/kg of HEI3090 or vehicle were given daily until day 21. (B) Representative images of lung sections at day 21 after treatment stained with H&E and Sirius Red, Scale bar= 100 µm. (C) Fibrosis score assessed by the Ashcroft method. P-values were determined by Two-tailed unpaired t-test. (D) Quantification of fibroblasts/myofibroblasts in non-immune cell subset. P-values were determined by Two-tailed Mann-Whitney test. (E) Representative images of lung sections at day 21 after treatment stained with Sirius Red, Scale bar= 100 µm. (F) Collagen levels in whole lung of mice assessed on Sirius Red-polarized images. P-values were determined by Two-tailed unpaired t-test. (G) Experimental design. WT mice were given 2.5 U/kg of bleomycin by i.n. route. 1.5 mg/kg of HEI3090 or vehicle were given daily until day 14. (H) Representative images of lung sections at day 14 after treatment stained with H&E, Scale bar= 100 µm. (I) Fibrosis score assessed by the Ashcroft method. P-values were determined by Two-tailed unpaired t-test. (J) Quantification of fibroblasts/myofibroblasts in non-immune cell subset. P-values were determined by Two-tailed Mann-Whitney test. (K) Representative images of lung sections at day 14 after treatment stained with Sirius Red, Scale bar= 100 µm. (L) Collagen levels in whole lung of mice assessed on Sirius Red-polarized images. P-values were determined by Two-tailed Mann-Whitney test. (M) P2rx7 KO mice were given 2.5 U/kg of bleomycin by i.n. route. 1.5 mg/kg of HEI3090 or vehicle were given daily until day 14. Fibrosis score assessed by the Ashcroft method is showed. P-values were determined by Two-tailed Mann-Whitney test. For all analyses, the violin plot illustrates the distribution of Ashcroft scores across indicated groups. The width of the violin at each point represents the density of data, and the central line indicates the median expression level. Each point represents one biological replicate. *p < 0.05, **p < 0.01. WT: Wildtype, BLM: bleomycin, i.p.: intraperitoneal, i.n.: intranasal, H&E: hematoxylin & eosin, AU: arbitrary units.

As accumulation of fibroblasts/myofibroblasts and extracellular matrix proteins are a hallmark of fibrosis, we also quantified fibroblasts/myofibroblasts and collagen levels in the lungs of vehicle and HEI3090-treated mice using respectively an antibody which is specific to the PDGFRα and by measuring the intensity of Sirius Red-stained collagen fibers under polarized light. We showed that both the percentage of CD140a (PDGFRα) positive cells in non-immune cells subset (Figure 2D) and the collagen fibers were reduced in lungs of HEI3090-treated mice (Figure 2, E and F).

Earlier version: We also tested the ability of HEI3090 to limit lung fibrosis progression when administered during the inflammatory phase of the BLM model (Figure 2G) that mimics the exacerbation episodes in IPF patients [44]. HEI3090 was also able to inhibit the onset of lung fibrosis in this setting given less damage on lung architecture (Figure 2H), the reduced fibrosis score (Figure 2I). We further evaluated the percentage of fibroblasts/myofibroblasts in non-immune cells subset (Figure 2J). At this shorter time point, we observed less fibroblasts/myofibroblasts in lungs of WT mice treated with vehicle and the number of CD140a positive cells is decreased in lungs of HEI3090-treated mice. As for mice with an established fibrosis, HEI3090 reduced the intensity of collagen fibers in lungs of HEI3090-treated mice (Figure 2, K and L). HEI3090 being an activator of P2RX7 [42], we verified that its antifibrotic action depends on the expression of P2RX7 using P2rx7 KO mice. We showed that HEI3090 is unable to limit lung fibrosis progression (Figure 2M, Supplementary Figure 2C and D). Importantly, a Dunnett’s multiple comparison test with data from Figure 2I and M rules out a genetic background effect of HEI3090 efficacy to attenuate lung fibrosis (Supplementary Figure 2E).

Revised version: We also tested the ability of HEI3090 to limit lung fibrosis progression when administered during the inflammatory phase of the BLM model (Figure 2G) that mimics the exacerbation episodes in IPF patients (Peng et al., 2013). HEI3090 was also able to inhibit the onset of lung fibrosis in this setting given less damage on lung architecture (Figure 2H), the reduced fibrosis score (Figure 2I). We further evaluated the percentage of fibroblasts/myofibroblasts in non-immune cells subset (Figure 2J). At this shorter time point, we observed less fibroblasts/myofibroblasts in lungs of WT mice treated with vehicle and the number of CD140a positive cells is decreased in lungs of HEI3090-treated mice. Further, and as for mice with an established fibrosis, HEI3090 reduced the intensity of collagen fibers in lungs of HEI3090-treated mice (Figure 2, K and L). We further verified whether the antifibrotic action of HEI3090 depends on the expression of P2RX7 by inducing lung fibrosis in p2rx7 knockout (KO) mice. In doing so, we initially observed that P2RX7 plays a role in the development of BLM-induced lung fibrosis. This is illustrated by a decrease of 50% in the Ashcroft score, with a mean value of 1.7 in P2RX7 knockout mice compared to 3 in wild-type mice (Figure 2M and Supplemental Figure 2C). It is important to note that p2rx7 -/-mice still exhibit signs of lung fibrosis, such as thickening of the alveolar wall and a reduction in free air space, in comparison to naïve mice that received PBS instead of BLM (see Supplemental Figure 2A). This result confirms a previous report indicating that BLM-induced lung fibrosis partially depends on the activation of the P2RX7/pannexin-1 axis, leading to the production of IL-1β in the lung. Additionally, in contrast to the observations in WT mice, HEI3090 failed to attenuate the remaining lung fibrosis in p2rx7 -/- mice, as measured by the Ashcroft score (Figure 2M), the percentage of lung tissue with fibrotic lesions, or the intensity of collagen fibers (Supplemental Figure 2D). These results show that P2RX7 alone participates in fibrosis and that HEI3090 exerts a specific antifibrotic effect through this receptor (see Supplemental Figure 2C).

Removed in the revised version: Hence, the three different modes of analysis show that the P2RX7 activator, HEI3090, inhibits the lung fibrosis progression and is effective during both the fibroproliferative and acute inflammation phase of the BLM-induced pulmonary fibrosis mouse model. Therefore, and to reduce the number of mice, we used in the rest of the study the Ashcroft score and collagen fiber measurement to characterize the antifibrotic effect of HEI3090.

HEI3090 shapes immune cell infiltration in the lungs

Earlier version: Since transcriptomic analysis showed that P2RX7 has immunoregulatory functions in IPF (Figure 1) and HEI3090 has antifibrotic activity (Figure 2), we next investigated if HEI3090 had an impact on both the immune landscape of the lung and production of cytokines. We show that lung CD3+ T cells were more biased to produce IFN-γ than the profibrotic IL-17A cytokine when mice were treated with HEI3090 (Figure 3A and Supplemental Figure 3A). This biased production of IFN-γ is only seen in CD3+ T cells and not in overall lung immune cells (Supplemental Figure 3B) nor in the subsets of T lymphocytes (Figure 3B, Supplemental Figure 3B) or NK cells (Supplemental Figure 3C). Although levels of CD3+ T cells and T cell subsets were unchanged (Supplemental Figure 3D), including the profibrotic Th17 cells (Figure 3C), IL-17A production by Th17 cells is substantially attenuated after HEI3090 treatment (Figure 3C), consistent with the ability of IFN-γ to inhibit IL-17A production [45]. Considering the strong profibrotic properties of TGFβ and its mutual antagonism with IFN-γ [46,47], we checked whether HEI3090 had an effect on TGFβ levels. Indeed, treatment with HEI3090 reduced TGFβ-producing immune cells in the lung as well as TGFβ production (Figure 3D). Notably, HEI3090 treatment reduced TGFβ production in NK cells but not in T-cell subsets (Supplemental Figure 3F).

Revised version: Given that transcriptomic analysis revealed the immunoregulatory functions of P2RX7 in IPF (Figure 1) and the antifibrotic activity of HEI3090 (Figure 2), we subsequently explored whether HEI3090 had an impact on both the immune landscape of the lung and the production of cytokines. This investigation aimed to elucidate the potential mechanisms underlying the dual antifibrotic effects observed with both P2RX7 deletion and P2RX7 activation by HEI3090. We show that lung CD3+ T cells were more biased to produce IFN-γ than the profibrotic IL-17A cytokine when mice expressing P2RX7 were treated with HEI3090 (Figure 3A and Supplemental Figure 3A). This biased production of IFN-γ is only seen in CD3+ T cells and not in overall lung immune cells (Supplemental Figure 3B) nor in the subsets of T lymphocytes (Figure 3B, Supplemental Figure 3B) or NK cells (Supplemental Figure 3C). Although levels of CD3+ T cells and T cell subsets were unchanged (Supplemental Figure 3D), including the profibrotic Th17 cells (Figure 3C), IL-17A production by Th17 cells is substantially attenuated after HEI3090 treatment (Figure 3C), consistent with the ability of IFN-γ to inhibit IL-17A production (Laurence et al., 2007). Considering the strong profibrotic properties of TGFβ and its mutual antagonism with IFN-γ (Ishida et al., 2004; Park et al., 2007), we checked whether HEI3090 had an effect on TGFβ levels. Indeed, treatment with HEI3090 reduced TGFβ-producing immune cells in the lung as well as TGFβ production (Figure 3D). Notably, HEI3090 treatment reduced TGFβ production in NK cells but not in T-cell subsets (Supplemental Figure 3F).

HEI3090 favors an anti-fibrotic immune signature in the lungs

Earlier version: (A) WT mice were given 2.5 U/kg of bleomycin by i.n. route and treated daily i.p. with 1.5 mg/kg of HEI3090 or Vehicle. Lungs were analyzed by flow cytometry at day 14. (A) Contour plot of IFN-γ and IL-17A producing T cells (CD3+NK1.1-) (left) and ratio of IFN-γ over IL- 17A in T cells (CD3+NK1.1-) (right). (B) Percentage of IFN-γ producing CD4+ and CD8+ T cells. (C) Percentage and GMFI of IL-17A+ cells of CD4+ T cells (CD3+CD4+NK1.1-). (D) Percentage and GMFI of TGF in CD45+ cells. (E) Dotplot showing lung inflammatory monocytes, gated on lineage-CD11c-CD11b+ cells (left) and percentage of lung inflammatory monocytes (Ly6ChighLy6G-) (right) (F) Percentage of alveolar macrophages (CD11c+SiglecF+) and (G) lung eosinophils (CD11b+SiglecF+CD11c-). Each point represents one mouse, data represented as violin plots or mean±SEM, two-tailed Mann-Whitney text, *p<0.05, **p<0.01. GMFI: geometric mean fluorescence intensity. i.n.: intranasal, i.p.: intraperitoneal.

Revised version: WT mice were given 2.5 U/kg of bleomycin by i.n. route and treated daily i.p. with 1.5 mg/kg of HEI3090 or Vehicle. Lungs were analyzed by flow cytometry at day 14. (A) Contour plot of IFN-γ and IL-17A producing T cells (CD3+NK1.1-) (left) and ratio of IFN-γ over IL-17A in T cells (CD3+NK1.1-) (right). P-values were determined by Two-tailed Mann-Whitney test. (B) Percentage of IFN-γ producing CD4+ and CD8+ T cells. P-values were determined by Two-tailed Mann-Whitney test. (C) Percentage and GMFI of IL-17A+ cells of CD4+ T cells (CD3+CD4+NK1.1-). P-values were determined by Two-tailed Mann-Whitney test. (D) Percentage and GMFI of TGFβ in CD45+ cells. P-values were determined by Two-tailed Mann-Whitney test. (E) Dotplot showing lung inflammatory monocytes, gated on lineage-CD11c- CD11b+ cells (left) and percentage of lung inflammatory monocytes (Ly6ChighLy6G-) (right). P-values were determined by Two-tailed Mann-Whitney test. (F) Percentage of alveolar macrophages (CD11c+SiglecF+), p-values were determined by Two-tailed unpaired t-test and (G) lung eosinophils (CD11b+SiglecF+CD11c-), p-values were determined by Two-tailed Mann-Whitney test. (A-G), The violin plot illustrates the distribution of Ashcroft scores across indicated groups. The width of the violin at each point represents the density of data, and the central line indicates the median expression level. Each point represents one biological replicate. *p < 0.05, **p < 0.01. GMFI: geometric mean fluorescence intensity. i.n.: intranasal, i.p.: intraperitoneal.

Earlier version: Pulmonary fibrosis is also favored and driven by the recruitment of inflammatory cells, mainly from the myeloid lineage. Monocytes are highly inflammatory cells that are recruited to the lung and differentiate into alveolar macrophages, both of which have strong profibrotic properties [48-50]. We demonstrated that in HEI3090-treated mice, the number of inflammatory monocytes decreased markedly (Figure 3G), whereas the number of alveolar macrophages remained unchanged (Figure 3F), consistent with the prognostic ability of monocyte count in IPF progression [51-54]. We also examined other inflammatory cells with profibrotic properties, such as eosinophils [55] that are less present in HEI3090-treated lungs (Figure 3G), or PMN levels that remain unchanged by HEI3090 treatment (Supplemental Figure 3D).

Revised version: Pulmonary fibrosis is also favored and driven by the recruitment of inflammatory cells, mainly from the myeloid lineage. Monocytes are highly inflammatory cells that are recruited to the lung and differentiate into alveolar macrophages, both of which have strong profibrotic properties (Gibbons et al., 2011; McCubbrey et al., 2018; Misharin et al., 2017). We demonstrated that in HEI3090-treated mice, the number of inflammatory monocytes decreased markedly (Figure 3G), whereas the number of alveolar macrophages remained unchanged (Figure 3F), consistent with the prognostic ability of monocyte count in IPF progression (Achaiah et al., 2021; Kreuter et al., 2021; Liu et al., 2019; Scott et al., 2019). We also examined other inflammatory cells with profibrotic properties, such as eosinophils (Morimoto et al., 2018) that are less present in HEI3090-treated lungs (Figure 3G), or PMN levels that remain unchanged by HEI3090 treatment (Supplemental Figure 3D).

We also wondered if the activation of P2RX7 has a systemic effect by analyzing immune changes in mice’s spleens. No significant change in cell populations was observed in the spleens of mice (Supplemental Figure 4B) when HEI3090 was administrated in the early phase of the BLM model suggesting a local lung activity of the molecule. However, HEI3090 reactivated a systemic immune response with higher levels of dendritic cells and lymphocytes in the spleens of mice treated during the fibroproliferation phase (Supplemental Figure 4D). These results show the ability of HEI3090 to shape the immune response locally and impact the progression of fibrosis systemically even in the fibroproliferative phase of the BLM model.

Earlier version: Altogether these results demonstrate that activation of P2RX7 with HEI3090 promotes an antifibrotic cytokinic profile in lung immune cells and attenuates lung inflammation.

Revised version: Altogether these results demonstrate that HEI3090 promotes an antifibrotic cytokinic profile in lung immune cells and attenuates lung inflammation.

HEI3090 requires the P2RX7/NLRP3/IL-18 pathway in immune cells to inhibit lung fibrosis

Removed in the revised version: We wanted to further investigate the mechanism of action of HEI3090 by identifying the cellular compartment and signaling pathway required for its activity. Since the expression of P2RX7 and the P2RX7-dependent release of IL-18 are mostly associated with immune cells[56], and since HEI3090 shapes the lung immune landscape (Figure 3), we investigated whether immune cells were required for the antifibrotic effect of HEI3090.

Earlier version: To do so, we performed an adoptive transfer experiment with WT P2RX7-expressing splenocytes (Figure 4A, Supplemental Figure 5E) into p2rx7-/- mice one day before BLM administration, this approach being previously used with success [42]. We show that restriction of P2RX7 expression on immune cells restored the antifibrotic effect of HEI3090 based on the architecture of the lung, with lungs of HEI3090-treated mice showing more free airspace and thinner alveolar walls (Figure 4B), as well as an overall lower fibrosis score (Figure 4C) and collagen fiber intensity (Figure 4D) than control lungs. Since the bleomycin mouse model relies on P2RX7-expressing epithelial cells [5], we wanted to validate further the role of P2RX7- expressing immune cells in a mouse model where P2RX7 is expressed by non-immune cells. To do so, we reduced both the expression of P2RX7 and its activity by repeated i.v. administration of p2rx7-/- splenocytes in WT mice (Supplemental Figure 5). In this setting, HEI3090 was unable to limit the progression of fibrosis. Moreover, we show that the activity of HEI3090 requires P2RX7 expression, as this effect was lost in p2rx7-/- mice (Figure 2M, and Supplemental Figures 2C and 6, A and B) [42]. These results highlight the important role of immune cells and rules out a major role of non-immune P2RX7-expressing cells, such as fibroblasts, in the antifibrotic effect of HEI3090.

Revised version: We wanted to further investigate the mechanism of action of HEI3090 by identifying the cellular compartment and signaling pathway required for its activity. Since the expression of P2RX7 and the P2RX7-dependent release of IL-18 are mostly associated with immune cells (Ferrari et al., 2006), and since HEI3090 shapes the lung immune landscape (Figure 3), we investigated whether immune cells were required for the antifibrotic effect of HEI3090. To do so, we conducted adoptive transfer experiments wherein immune cells from a donor mouse were intravenously injected one day before BLM administration into an acceptor mouse. The intravenous injection route was chosen as it is a standard method for targeting the lungs, as previously documented (Wei and Zhao, 2014). This approach was previously used with success in our laboratory (Douguet et al., 2021). It is noteworthy that this adoptive transfer approach did not influence the response to HEI3090. This was observed consistently in both p2rx7 -/- mice and p2rx7 -/- mice that received splenocytes of the same genetic background. In both cases, HEI3090 failed to mitigate lung fibrosis, as depicted in Figure 2M and Supplemental Figures 2D and 6A and B. First, we tested whether WT P2RX7-expressing splenocytes infuse into p2rx7 -/- mice restored the antifibrotic activity of HEI3090. We show that restriction of P2RX7 expression on immune cells restored the antifibrotic effect of HEI3090 based on the architecture of the lung, with lungs of HEI3090-treated mice showing more free airspace and thinner alveolar walls (Figure 4B), as well as an overall lower fibrosis score (Figure 4C) and collagen fiber intensity (Figure 4D) than control lungs. Since it was previously showed that the bleomycin mouse model relies on P2RX7-expressing pulmonary epithelial cells (Riteau et al., 2010), we wanted to validate further the role of P2RX7-expressing immune cells in a mouse model where P2RX7 is expressed by non-immune cells. To do so, we reduced both the expression of P2RX7 and its activity by repeated i.v. administration (days 1, 4, 7 9 and 11) of p2rx7-/- splenocytes in WT mice (Supplemental Figure 5). In this setting, HEI3090 was unable to limit the progression of fibrosis. These results highlight the important role of immune cells and rules out a major role of non-immune P2RX7-expressing cells, such as fibroblasts, in the antifibrotic effect of HEI3090.

The P2RX7/NLRP3/IL-18 pathway in immune cells is required for HEI3090’s antifibrotic effect

Earlier version: (A) Experimental design. p2rx7-/- mice were given 3.106 WT, nlrp3-/- or il18-/- splenocytes i.v. one day prior to BLM delivery (i.n. 2.5 U/kg). Mice were treated daily i.p. with 1.5 mg/kg HEI3090 or vehicle for 14 days. (B,E,H) Representative images of lung sections at day 14 after treatment stained with H&E and Sirius Red, scale bar= 100 µm. Fibrosis score assessed by the Ashcroft method of adoptive transfer of WT splenocytes, (C) nlrp3-/- splenocytes (F) and (I) il18-/- splenocytes. Collagen fiber intensity of adoptive transfer of WT splenocytes, (D) nlrp3-/- splenocytes (G) and (J) il18-/- splenocytes. (K) IL-18 levels in sera of WT BLM-induced mice at day 14 (L) Ratio of IFN-γ over IL-17A in lung T cells (CD3+NK1.1-) of WT mice neutralized for IL-18 or not (isotype control) every 3 days. Each point represents one mouse, two-tailed Mann-Whitney test, *p<0.0.5. WT: Wildtype, BLM: bleomycin, i.p: intraperitoneal, i.n.: intranasal, i.v.: intravenous, H&E: hematoxylin & eosin.

Revised version: (A) Experimental design. p2rx7-/- mice were given 3.106 WT, nlrp3-/- or il18-/- splenocytes i.v. one day prior to BLM delivery (i.n. 2.5 U/kg). Mice were treated daily i.p. with 1.5 mg/kg HEI3090 or vehicle for 14 days. (B,E,H) Representative images of lung sections at day 14 after treatment stained with H&E and Sirius Red, scale bar= 100 µm. (C, F, I), Fibrosis score assessed by the Ashcroft method and (D, G, J), Collagen fiber intensity. (C) WT splenocytes, p-values were determined by Two-tailed Mann-Whitney test. (D) WT splenocytes, p-values were determined by Two-tailed unpaired t-test. (F, G) nlrp3-/- splenocytes, p-values were determined by Two-tailed Mann-Whitney test. (I, J) il18-/- splenocytes, p-values were determined by Two-tailed Mann-Whitney test. (K) IL-18 levels in sera of WT BLM-induced mice at day 14. P-values were determined by Two-tailed unpaired t-test. (L) Ratio of IFN-γ over IL-17A in lung T cells (CD3+NK1.1-) of WT mice neutralized for IL-18 or not (isotype control) every 3 days. P-values were determined by one-way ANOVA with tukey’s multiple comparisons test. For all analyses, the violin plot illustrates the distribution of Ashcroft scores across indicated groups. The width of the violin at each point represents the density of data, and the central line indicates the median expression level. Each point represents one biological replicate. *p < 0.05, **p < 0.01 WT: Wildtype, BLM: bleomycin, i.p.: intraperitoneal, i.n.: intranasal, i.v.: intravenous, H&E: hematoxylin & eosin.

Earlier version: To test the importance of the NLRP3/IL-18 pathway downstream of P2RX7, we performed an adoptive transfer of nlrp3-/- and il18-/- splenocytes into p2rx7-/- mice, expressing similar levels of P2RX7 as WT splenocytes (Supplemental Figure 6E), but also the same levels of IL-18 and NLRP3 (Supplemental Figure 6F). The absence of NLRP3 and IL-18 in P2RX7-expressing immune cells abrogated the ability of HEI3090 to inhibit lung fibrosis because the lung architecture resembled that of control mice as well as the collagen fiber intensity (Figure 4, E-J). Consistent with the requirement of IL- 8 for HEI3090’s antifibrotic activity, activation of P2RX7 with this molecule in WT mice increased the levels of IL-18 in the sera of these mice compared to control mice (Figure 4K). Moreover, neutralization of IL-18 abrogated the increase of the IFN-γ/IL-17A ratio by lung T cells (Figure 4L), highlighting furthermore the necessity of IL-18 for the antifibrotic effect of HEI3090.

Revised version: To test the importance of the NLRP3/IL-18 pathway downstream of P2RX7, we performed an adoptive transfer of nlrp3-/- and il18-/- splenocytes into p2rx7-/-mice, expressing similar levels of P2RX7 as WT splenocytes (Supplemental Figure 6E), but also the same levels of IL-18 and NLRP3 (Supplemental Figure 6F). We also verified in this experimental setting by comparing each experimental groups (treated and non-treated) that the genetic background did not affect lung fibrosis (Supplementary Figure 7). The absence of NLRP3 and IL-18 in P2RX7-expressing immune cells abrogated the ability of HEI3090 to inhibit lung fibrosis because the lung architecture and collagen fibers intensity resembled that of control mice (Figure 4, E-J). Consistent with the requirement of IL-18 for HEI3090’s antifibrotic activity, activation of P2RX7 with this molecule in WT mice increased the levels of IL-18 in the sera of these mice compared to control mice (Figure 4K). Moreover, neutralization of IL-18 abrogated the increase of the IFN-γ/IL-17A ratio by lung T cells (Figure 4L), highlighting furthermore the necessity of IL-18 for the antifibrotic effect of HEI3090.

Earlier version: Not only does the activation of the P2RX7/NLRP3 pathway leads to the release of IL-18, but also induce the release of the pro-inflammatory and pro-fibrotic IL-1β cytokine. However, IL-1β was not involved in the antifibrotic effect of HEI3090 (Supplemental Figure 6C), nor were its levels affected by HEI3090 in WT mice (Supplemental Figure 6D).

Revised version: Not only does the activation of the P2RX7/NLRP3 pathway leads to the release of IL-18, but also induce the release of the pro-inflammatory and pro-fibrotic IL-1β cytokine. However, IL-1β was not involved in the antifibrotic effect of HEI3090 (Supplemental Figure 6C), nor were its levels affected by HEI3090 in WT mice (Supplemental Figure 6D), in contrast to IL-18 levels measured in the same biological sample (Figure 4K).

Overall, we show that the P2RX7/NLRP3/IL-18 axis in immune cells is required to limit lung fibrosis progression, highlighting the efficacy in targeting the immune system in this disease.

Discussion

Added in the revised version: A major unmet need in the field of IPF is new treatment to fight this uncurable disease. In this preclinical study, we demonstrate the ability of immune cells to limit lung fibrosis progression. Based on the hypothesis that a local activation of a T cell immune response and upregulation of IFN-γ production has antifibrotic proprieties, we used the HEI3090 positive modulator of the purinergic receptor P2RX7, previously developed in our laboratory (Douguet et al., 2021), to demonstrate that activation of the P2RX7/IL-18 pathway attenuates lung fibrosis in the bleomycin mouse model. We have demonstrated that lung fibrosis progression is inhibited by HEI3090 in the fibrotic phase but also in the acute phase of the BLM fibrosis mouse model, i.e. during the period of inflammation. This lung fibrosis mouse model commonly employed in preclinical investigations, has recently been recognized as the optimal model for studying IPF (Jenkins et al., 2017). In this model, the intrapulmonary administration of BLM induces DNA damage in alveolar epithelial type 1 cells, triggering cellular demise and the release of ATP. The extracellular release of ATP from injured cells activates the P2RX7/pannexin 1 axis, initiating the maturation of IL1β and subsequent induction of inflammation and fibrosis. In line with this, mice lacking P2RX7 exhibited reduced neutrophil counts in their bronchoalveolar fluids and decreased levels of IL1β in their lungs compared to WT mice (Riteau et al., 2010). Based on these findings, Riteau and colleagues postulated that the inhibition of P2RX7 activity may offer a potential strategy for the therapeutic control of fibrosis in lung injury. In the present study we provided strong evidence showing that selective activation of P2RX7 on immune cells, through the use of HEI3090, can dampen inflammation and fibrosis by releasing IL-18. The efficacy of HEI3090 to inhibit lung fibrosis was evaluated histologically on the whole lung’s surface by evaluating the severity of fibrosis using three independent approaches applied to the whole lung, the Ashcroft score, quantification of fibroblasts/myofibroblasts (CD140a) and polarized-light microscopy of Sirius Red staining to quantify collagen fibers. All these methods of fibrosis assessment revealed that HEI3090 exerts an inhibitory effect on lung fibrosis, underscoring the necessity for a thorough pre-clinical assessment of HEI3090’s mode of action. Notably, HEI3090 functions as an activator, rather than an inhibitor, of P2RX7, further emphasizing the importance of elucidating its intricate mechanisms.

Removed in the revised version: A major unmet need in the field of IPF is new treatment to fight this uncurable disease. In this preclinical study, we demonstrate the ability of immune cells to limit lung fibrosis progression. Based on the hypothesis that a local activation of a T cell immune response and upregulation of IFN-γ production has antifibrotic proprieties, we used the HEI3090 positive modulator of the purinergic receptor P2RX7, previously developed in our laboratory [42], to demonstrate that activation of the P2RX7/IL-18 pathway attenuates lung fibrosis in the bleomycin mouse model.

Removed in the revised version: We have demonstrated that lung fibrosis progression is inhibited by HEI3090 in the fibrotic phase but also in the acute phase of the BLM fibrosis mouse model, i.e. during the period of inflammation. This lung fibrosis mouse model is classically used in preclinical studies and has been designated recently as the best model for IPF [57]. The efficacy of HEI3090 to inhibit lung fibrosis was evaluated histologically on the whole lung’s surface by evaluating the severity of fibrosis using three independent approaches applied to the whole lung, the Ashcroft score (52), quantification of fibroblasts/myofibroblasts (CD140a) and polarized-light microscopy of Sirius Red staining to quantify collagen fibers. In all settings, HEI3090 reduced alveolar wall thickness, the percentage of CD140a positive cells in the non-immune cell subset and accumulation of collagen fibers, highlighting the need of a comprehensive pre-clinical evaluation of the mode of action of HEI3090 to reduce lung fibrosis.

Earlier version: Our study showed that inhibition of fibrosis progression by HEI3090 was associated with an increased production of IFN-γ by lung T cells that was dependent on an increased release of IL-18. We also showed that expression of the P2RX7/IL-18/IFN-γ pathway is attenuated in IPF patients where TGFβ levels are high [58], consistent with the ability of TGFβ to downregulate IL-18R expression and IL-18-mediated IFN-γ production [59]. These results confirm the beneficial effects of enhancing activation of the P2RX7/IL-18/IFN-γ pathway.

Revised version: Our study showed that inhibition of fibrosis by HEI3090 was associated with an increased production of IFN-γ by lung T cells that was dependent on an increased release of IL-18. The transcriptional analysis additionally indicated the downregulation of the P2RX7/IL-18/IFN-γ pathway in patients with IPF, where TGFβ levels are high (Khalil et al., 1991). This is consistent with the ability of TGFβ to downregulate IL-18R expression and IL-18-mediated IFN-γ production (Koutoulaki et al., 2010). These results confirm the beneficial effects of enhancing activation of the P2RX7/IL-18/IFN-γ pathway.

Earlier version: P2RX7 is expressed by various immune and non-immune cells, but its expression is the highest in dendritic cells (DCs) and macrophages [60], from which IL-18 is mainly released to shape the T-cell response [61] and increase T-cell IFN-γ production [62]. Collectively, and consistent with the immunomodulatory properties of P2RX7, these observations suggest that HEI3090 may target P2RX7-expressing antigen-presenting cells to influence the T-cell response, which may explain the selective T-cell increase in IFN-γ in HEI3090-treated mice. Accordingly, we have previously shown that HEI3090 targets the P2RX7/IL-18 axis in DCs to shape the immune response in a lung tumor mouse model [42]. Noteworthy, by contrast to parenteral administration of IFN-γ which failed in clinical trial [31], we did not observe any sign of systemic side effect likely because HEI3090 needs ATP to potentiate P2RX7 and ATP is present only in injured tissues.

Revised version: P2RX7 is expressed by various immune and non-immune cells, but its expression is the highest in dendritic cells (DCs) and macrophages (Di Virgilio and Vuerich, 2015), from which IL-18 is mainly released to shape the T-cell response (Sáez et al., 2017) and increase T-cell IFN-γ production (Iwai et al., 2008). Collectively, and consistent with the immunomodulatory properties of P2RX7, these observations suggest that HEI3090 may target P2RX7-expressing antigen-presenting cells to influence the T-cell response, which may explain the selective T-cell increase in IFN-γ in HEI3090-treated mice. Accordingly, we have previously shown that HEI3090 targets the P2RX7/IL-18 axis in DCs to shape the immune response in a lung tumor mouse model (Douguet et al., 2021). Noteworthy, by contrast to parenteral administration of IFN-γ which failed in clinical trial (King et al., 2009), we did not observe any sign of systemic side effect likely because HEI3090 needs ATP to potentiate P2RX7 and ATP is present only in injured tissues.

Earlier version: Activation of P2RX7 with HEI3090 not only increased IFN-γ production by T cells but it also reshaped the immune and cytokinic composition of the lung. Indeed, lungs of HEI3090-treated mice show a decrease in IL-17A production by Th17 cells and TGFβ production by lung immune cells. Moreover, lung inflammation is dampened after HEI3090 treatment, since the number of inflammatory monocytes and eosinophils decreases. It is not known whether this cytokinic and anti-inflammatory switch is solely due to the IL-17A and TGFβ-suppressive property of IFN-γ [63-65], or whether it is a combination with the cell death-inducing property of P2RX7 [6].

Revised version: HEI3090 not only increased IFN-γ production by T cells but it also reshaped the immune and cytokinic composition of the lung. Indeed, lungs of HEI3090-treated mice show a decrease in IL-17A production by Th17 cells and TGFβ production by lung immune cells. Moreover, lung inflammation is dampened after HEI3090 treatment, since the number of inflammatory monocytes and eosinophils decreases. It is not known whether this cytokinic and anti-inflammatory switch is solely due to the IL-17A and TGFβ-suppressive property of IFN-γ (de Bruin et al., 2010; Iwamoto et al., 1993; Penton-Rol et al., 1998), or whether it is a combination with the cell death-inducing property of P2RX7 (Surprenant et al., 1996).

Removed in the revised version: In addition to HEI3090, other molecules have been suggested to be positive modulators of P2RX7 in vitro [66]. Among them the most described are clemastine and ginsenosides. Clemastine, a first generation anti-histaminic, enhances ATP-induced P2RX7 dependent Ca2+ entry, macropore opening and IL1B production in immune cells [67]. However, when tested in vivo clemastine treatment does not activates P2RX7, indeed it reduces the production of pro inflammatory cytokines and decreases P2RX7 expression in hippocampus [68]. These biological responses correspond to an inhibitory effect on P2RX7. Ginsenosides, purified from root of Panax ginseng, have been shown to potentiate ATP-induced P2RX7 dependent Ca2+ entry and macropore opening but also cell death in vitro [69]. So far, there is no evidence indicating that ginsenosides act as positive modulators of P2RX7 in vivo and all the other chemical compounds described to be active in vivo are negative allosteric modulator of P2RX7 [70].

Removed in the revised version: The novelty of our approach is that it targets and alters the immune environment of the lung. Indeed, the use of a P2RX7-specific modulator that acts only in an ATP-rich environment was effective in promoting an anti-inflammatory and anti-fibrotic phenotype by altering several key mediators of the disease. Therapies given to IPF patients are known to induce significant side effects which can limit their prescription [71]. Since current therapies (pirfenidone or nintedanib) and HEI3090 have different mechanisms of action, it could be interesting to test in preclinical studies the efficacy of a treatment combining lower doses of current molecules and HEI3090. If successful, this may open a new therapeutic option to be tested in IPF patients.

Earlier version: Of note, P2RX7 being known to be pro-inflammatory, through IL-1β-release, our strategy to activate P2RX7 with HEI3090 may be considered as risky. However, and consistent with our previous in vivo and in vitro studies [42] HEI3090 was unable to increase the release of IL-1β in the serum of BLM-treated mice even though it efficiently increased IL-18 release. Interestingly, it has been reported that ATP stimulation of alveolar macrophages derived from both IPF, and lung cancer patients resulted only in an increase in IL-18 [72,73], which was explained by an impaired NLRP3 inflammasome and a defective autophagy in IPF patients [74]. Autophagy can be regulated by P2RX7 (64). Indeed, whereas basal autophagy suppresses NLRP3- inflammasome activation and cytokine secretion, induced-autophagy augments IL-1β secretion [76], therefore ruling out the hypothesis of HEI3090-induced autophagy to control the production of IL1β. Unlike IL-1β, IL-18 is constitutively expressed in human and mouse immune cells [43] but also in non-immune cells such as fibroblasts [77] and lung epithelium [78] and is both matured and released after NLRP3 activation. Therefore, it is currently not known whether the lack of IL-1β release is due to the different cytokine expression pattern or whether there is IL-1β-specific regulation following an enhanced activation of P2RX7 and additional studies are required to answer this question.

Revised version: Of note, P2RX7 being known to be pro-inflammatory, through IL-1β-release, our strategy to activate P2RX7 with HEI3090 may be considered as risky. However, in line with our prior in vivo and in vitro investigations (Douguet et al., 2021), HEI3090 failed to augment the release of IL-1β in the serum of mice treated with BLM, despite effectively increasing the release of IL-18. This observation elucidates why HEI3090 is effective in alleviating lung fibrosis. In agreement, previous reports have indicated that ATP stimulation of alveolar macrophages, derived from both IPF and lung cancer patients, led exclusively to an elevation in IL-18 levels (Lasithiotaki et al., 2018, 2016), which was explained by an impaired NLRP3 inflammasome and a defective autophagy in IPF patients (Patel et al., 2012). Autophagy can be regulated by P2RX7 (de Bruin et al., 2010). Indeed, whereas basal autophagy suppresses NLRP3-inflammasome activation and cytokine secretion, induced-autophagy augments IL-1β secretion (Dupont et al., 2011), therefore ruling out the hypothesis of HEI3090-induced autophagy to control the production of IL-1β. Unlike IL-1β, IL-18 is constitutively expressed in human and mouse immune cells (Puren et al., 1999) but also in non-immune cells such as fibroblasts (Artlett et al., 2011) and lung epithelium (Watanabe et al., 2009) and is both matured and released after NLRP3 activation. Therefore, it is currently not known whether the lack of IL-1β release is due to the different cytokine expression pattern or whether there is IL-1β-specific regulation following an enhanced activation of P2RX7 and additional studies are required to answer this question.

Added in the revised version: In addition to HEI3090, other molecules have been suggested to be positive modulators of P2RX7 in vitro (Stokes et al., 2020). Among them the most described are clemastine and ginsenosides. Clemastine, a first generation anti-histaminic, enhances ATP-induced P2RX7 dependent Ca2+ entry, macropore opening and IL1B production in immune cells (Nörenberg et al., 2011). However, when tested in vivo clemastine treatment does not activates P2RX7, indeed it reduces the production of pro inflammatory cytokines and decreases P2RX7 expression in hippocampus (Su et al., 2018). These biological responses correspond to an inhibitory effect on P2RX7. Ginsenosides, purified from root of Panax ginseng, have been shown to potentiate ATP-induced P2RX7 dependent Ca2+ entry and macropore opening but also cell death in vitro (Helliwell et al., 2015). So far, there is no evidence indicating that ginsenosides act as positive modulators of P2RX7 in vivo and all the other chemical compounds described to be active in vivo are negative allosteric modulator of P2RX7 (Karasawa and Kawate, 2016).

Added in the revised version: The novelty of our approach is that it targets and alters the immune environment of the lung. Indeed, the use of a P2RX7-specific modulator that acts only in an ATP-rich environment was effective in promoting an anti-inflammatory and anti-fibrotic phenotype by altering several key mediators of the disease. Therapies given to IPF patients are known to induce significant side effects which can limit their prescription (Raghu et al., 2015). Since current therapies (pirfenidone or nintedanib) and HEI3090 have different mechanisms of action, it could be interesting to test in preclinical studies the efficacy of a treatment combining lower doses of current molecules and HEI3090. If successful, this may open a new therapeutic option to be tested in IPF patients.

Earlier version: In this study, we emphasize the importance of IL-18 for an antifibrotic effect. Several studies have indicated that P2RX7/NLRP3/IL-18 promote disease progression using knock out mice or inhibitors [5,20,77]. However, experimental mouse models rely on lung epithelial cell injury that has been shown to activate the NLRP3 inflammasome in the lung epithelium as a first step [5,18] and release danger signals that activate the immune system as a second step. Therefore, initial lung injury to epithelial cells is reduced or absent in p2rx7-/- and nlrp3-/- mice, indicating that P2RX7 and NLRP3 are required for the establishment of the bleomycin mouse model rather than their role in an already established fibrosis, which has not yet been studied. In addition, NLRP3 and the release of IL-1β and IL-18 from fibroblasts have been shown to promote myofibroblast differentiation and extracellular matrix production [77,79,80]. These observations suggest that fibrosis mouse models initially rely on NLRP3 activation by nonimmune cells and encourage further studies on the contribution of the NLRP3 inflammasome in immune cells to fibrosis progression in vivo. Since we have highlighted the importance of this pathway in immune cells for delaying fibrosis progression, we propose that IL-18 may have different effects depending on the cell type.

Revised version: In this study, we emphasize the importance of IL-18 for an antifibrotic effect. Several studies have indicated that P2RX7/NLRP3/IL-18 promote disease progression using knock out mice or inhibitors (Artlett et al., 2011; Riteau et al., 2010; Zhang et al., 2019). However, experimental mouse models rely on lung epithelial cell injury that has been shown to activate the NLRP3 inflammasome in the lung epithelium as a first step (Gasse et al., 2009; Riteau et al., 2010) and release danger signals that activate the immune system as a second step. Therefore, initial lung injury to epithelial cells is reduced or absent in p2rx7-/- and nlrp3-/- mice, indicating that P2RX7 and NLRP3 are required for the establishment of the bleomycin mouse model rather than their role in an already established fibrosis, which has not yet been studied. In addition, NLRP3 and the release of IL-1β and IL-18 from fibroblasts have been shown to promote myofibroblast differentiation and extracellular matrix production (Artlett et al., 2011; Pinar et al., 2020; Watanabe et al., 2009). These observations suggest that fibrosis mouse models initially rely on NLRP3 activation by nonimmune cells and encourage further studies on the contribution of the NLRP3 inflammasome in immune cells to fibrosis progression in vivo. Since we have highlighted the importance of this pathway in immune cells for delaying fibrosis progression, we propose that IL-18 may have different effects depending on the cell type. This hypothesis could be substantiated through experimentation utilizing inducible NLRP3 and P2RX7 knockout mice, wherein the invalidation of NLRP3 and P2RX7 would be induced after the establishment of fibrosis.

Earlier version: Beside an urgent need of new treatments for IPF, there is also a lack of biomarkers, such as prognostic biomarkers, markers of disease activity, or markers of drug efficacy. Our results suggest the possible benefit of an active IL-18 in the pathophysiology of pulmonary fibrosis and warrant analysis of IL-18 as a promising biomarker for predicting outcome in IPF patients. Given the potential effects of pirfenidone and nintedanib on IL-18 levels in preclinical models [81-83], determining IL-18 shifts during treatment would be highly interesting to evaluate potential changes in patients’ outcome and to examine IL-18 levels which may be helpful in the long run for patient treatment strategy and subsequent introduction of pipeline drugs [84].

Revised version: Beside an urgent need of new treatments for IPF, there is also a lack of biomarkers, such as prognostic biomarkers, markers of disease activity, or markers of drug efficacy. Our results suggest the possible benefit of an active IL-18 in the pathophysiology of pulmonary fibrosis and warrant analysis of IL-18 as a promising biomarker for predicting outcome in IPF patients. Given the potential effects of pirfenidone and nintedanib on IL-18 levels in preclinical models (Mavrogiannis et al., 2022; Oku et al., 2008; Sharawy and Serrya, 2020), determining IL-18 shifts during treatment would be highly interesting to evaluate potential changes in patients’ outcome and to examine IL-18 levels which may be helpful in the long run for patient treatment strategy and subsequent introduction of pipeline drugs (Spagnolo et al., 2021).

Overall, we highlight in this study the ability of the P2RX7/NLRP3/IL-18 pathway in immune cells to inhibit the onset of lung fibrosis by using a positive modulator of P2RX7 that acts selectively in an eATP-rich environment such as fibrotic lung. The unique feature of this strategy is that it enhances the antifibrotic and it attenuates the pro-fibrotic properties of immune cells, with no reported side effects, at least in mice.

Methods

Data availability

In silico data used in this study were retrieved from Gene Expression Omnibus (GEO). We used the Gene expression profiling of “chronic lung disease for the Lung Genomics Research Consortium cohort” (GSE47460, GPL14550). 122 patients with usual Interstitial Pneumonitis (UIP)/ Idiopathic Pulmonary Fibrosis (IPF) and 91 healthy controls were analyzed in this study. Microarray was done on whole lung homogenate from these subjects. Expression profile belong to the Lung Tissue Research Consortium (LTRC).

Mice

Earlier version: Mice were housed under standardized light—dark cycles in a temperature-controlled airconditioned environment under specific pathogen-free conditions at IRCAN, Nice, France, with free access to food and water. All mouse studies were approved by the committee for Research and Ethics of the local authorities (CIEPAL #598, protocol number APAFIS 21052- 2019060610506376) and followed the European directive 2010/63/UE, in agreement with the ARRIVE guidelines. Experiments were performed in accord with animal protection representative at IRCAN. p2rx7-/- (B6.129P2-P2rx7tm1Gab/J, from the Jackson Laboratory) were backcrossed with C57BL/6J OlaHsD mice. C57BL/6J OlaHsD male mice (WT) were supplied from Envigo (Gannat, France). To define the size of experimental groups we did a pilot experiment, with 4 mice in each group, and a statistical forecasting study. The results indicate that 6 mice per group will give significant results. Being aware that mice can unexpectedly dye due to BLM treatment, we chose to include at least 8 mice per group. In all experiment, the presence of outlier was checked thanks to the ROULT method, and all data are kept in the analysis.

Revised version: Mice were housed under standardized light–dark cycles in a temperature-controlled air-conditioned environment under specific pathogen-free conditions at IRCAN, Nice, France, with free access to food and water. All mouse studies were approved by the committee for Research and Ethics of the local authorities (CIEPAL #598, protocol number APAFIS 21052-2019060610506376) and followed the European directive 2010/63/UE, in agreement with the ARRIVE guidelines. Experiments were performed in accord with animal protection representative at IRCAN. p2rx7−/− (B6.129P2-P2rx7tm1Gab/J, from the Jackson Laboratory) were backcrossed with C57BL/6J OlaHsD mice. C57BL/6J OlaHsD male mice (WT) were supplied from Envigo (Gannat, France). The determination of experimental group sizes involved conducting a pilot experiment with four mice in each group. Subsequently, a power analysis, based on the pilot experiment’s findings (which revealed a 40% difference with a standard error of 0.9, α risk of 0.05, and power of 0.8), was performed to ascertain the appropriate group size for studying the effects of HEI3090 on BLM-induced lung fibrosis. The results of the pilot experiment and power analysis indicated that a group size of four mice was sufficient to characterize the observed effects. For each full-scale experiment, we initiated the study with 6 to 8 mice per group, ensuring a minimum of 5 mice in each group for robust statistical analysis. Additionally, we systematically employed the ROULT method to identify and subsequently exclude any outliers present in each experiment before conducting statistical analyses.

Induction of lung fibrosis

Earlier version: WT or p2rx7-/- male mice (8 weeks) were anesthetized with ketamine (25mg/kg) and xylazine (2.5 mg/kg) under light isoflurane and were given 2.5 U/kg of bleomycin sulfate (Sigma-Aldrich) by intranasal route. Mice were treated i.p. every day with vehicle (PBS, 10% DMSO) or with HEI3090 (1.5 mg/kg in PBS, 10% DMSO) [42] starting D1 or D7 post bleomycin delivery, as mentioned in the figures. After 14 days of treatment, lungs were either fixed for paraffin embedding or weighted and analyzed by flow cytometry. When mentioned, 200 µg of anti-IL-18 neutralizing antibody (BioXcell) or isotype control (IgG2a, BioXcell) were given by i.p. every three days starting one day prior to BLM administration.

Revised version: WT or p2rx7-/- male mice (8 weeks) were anesthetized with ketamine (25mg/kg) and xylazine (2.5 mg/kg) under light isoflurane and were given 2.5 U/kg of bleomycin sulfate (Sigma-Aldrich) by intranasal route. Mice were treated i.p. every day with vehicle (PBS, 10% DMSO) or with HEI3090 (1.5 mg/kg in PBS, 10% DMSO) starting D1 or D7 post bleomycin delivery, as mentioned in the figures. After 14 days of treatment, lungs were either fixed for paraffin embedding or weighted and analyzed by flow cytometry. When mentioned, 200 µg of anti-IL-18 neutralizing antibody (BioXcell) or isotype control (IgG2a, BioXcell) were given by i.p. every three days starting one day prior to BLM administration.

Adoptive transfer in p2rx7 deficient mice

Earlier version: Spleens from C57BL/6J male mice (8-10 weeks) were collected and digested with the spleen dissociation kit (Miltenyi Biotech) according to the supplier’s instructions. 3.10° splenocytes were injected i.v. in p2rx7-/- mice 1 day before intranasal delivery of bleomycin. Mice were treated i.p. every day for 14 days with vehicle (PBS, 10% DMSO) or with HEI3090 (1.5 mg/kg in PBS, 10% DMSO). Nlrp3-/- spleens were a kind gift from Dr Laurent Boyer, il18-/- spleens from Dr George Birchenough and il1b-/- spleens from Dr Bernhard Ryffel, all on a CS7BL/6J background.

Revised version: Spleens from C57BL/6J male mice (8-10 weeks) were collected and digested with the spleen dissociation kit (Miltenyi Biotech) according to the supplier’s instructions. 3.106 splenocytes were injected i.v. in p2rx7-/-mice 1 day before intranasal delivery of bleomycin or when indicated in repeated administration. Mice were treated i.p. every day for 14 days with vehicle (PBS, 10% DMSO) or with HEI3090 (1.5 mg/kg in PBS, 10% DMSO). Nlrp3-/- spleens were a kind gift from Dr Laurent Boyer, il18-/- spleens from Dr George Birchenough and il1b-/- spleens from Dr Bernhard Ryffel, all on a C57BL/6J background.

Histology

Earlier version: Lungs were collected and fixed in 3% formamide for 16h prior inclusion in paraffin. Deparaffinized lungs sections (3 µm) were rehydrated and sequentially incubated with Hematoxylin, Mayer’s, Bluing Reagent and Eosin Y Solution as indicated by the manufacturer (H&E staining, Abcam, ab245880) or Picro Sirius Red Solution followed by Acetic Acid as indicated by the manufacturer (Picro Sirius Red staining, Abcam, ab150681). The severity of fibrosis was assessed on whole lungs using the Ashcroft modified method which assigns 8 grades (0: normal lung, 1: isolated alveolar septa with gentle fibrotic changes, 2: fibrotic changes of alveolar septa with knot-like formation, 3: contiguous fibrotic walls of alveolar septa, 4: single fibrotic masses, 5: confluent fibrotic masses, 6: large contiguous fibrotic masses, 7: air bubbles and 8: fibrous obliteration to quantify lung fibrosis) to reliable and reproducible quantify fibrosis-induced lung remodeling [85]. To be even more accurate and not biased by patchy lesions observed in all existing lung fibrosis induced mouse models, the whole lungs (left and right lobes) were divided in section of 880 µm2 and each section was scored individually as shown in Supplementary Figure 2. A total of 80 to 110 sections were analyzed per mouse.

Revised version: Lungs were collected and fixed in 3% formamide for 16 h prior inclusion in paraffin. Deparaffinized lungs sections (3 µm) were rehydrated and sequentially incubated with Hematoxylin, Mayer’s, Bluing Reagent and Eosin Y Solution as indicated by the manufacturer (H&E staining, Abcam, ab245880) or Picro Sirius Red Solution followed by Acetic Acid as indicated by the manufacturer (Picro Sirius Red staining, Abcam, ab150681). The severity of fibrosis was assessed on whole lungs using the Ashcroft modified method which assigns 8 grades (0: normal lung, 1: isolated alveolar septa with gentle fibrotic changes, 2: fibrotic changes of alveolar septa with knot-like formation, 3: contiguous fibrotic walls of alveolar septa, 4: single fibrotic masses, 5: confluent fibrotic masses, 6: large contiguous fibrotic masses, 7: air bubbles and 8: fibrous obliteration to quantify lung fibrosis) to reliable and reproducible quantify fibrosis-induced lung remodeling (Hübner et al., 2008). To be even more accurate and not biased by patchy lesions observed in all existing lung fibrosis induced mouse models, the whole lungs (left and right lobes) were divided in section of 880 µm2 and each section was scored individually as shown in Supplementary Figure 2. A total of 80 to 110 sections were analyzed per mouse.

Levels of collagen on whole lungs were assessed on Sirius Red polarized light images of the entire lung taken with HD - Axio Observer Z1 Microscope ZEISS microscope (Figure 2) or with the SLIDEVIEW V200 slide scanner, Evident Europe Rungis (Figure 4). The collagen amount given by the polarization intensity of the Sirius red staining of the lung slices and the percentage of fibrotic area was quantified with an ImageJ/Fiji macro program as following: The mean gray value and the integrated intensity of the collagen staining were measured in the fibrotic and non-fibrotic regions excluding the signal coming from vessels and bronchial tubes using dedicated masks. The binary masks were obtained after a median filtering and a manual thresholding, from the brightfield images for the fibrotic and non-fibrotic ones and the polarization images for the vessels and bronchial tubes. Vessels and bronchial tubes are excluded from fibrotic and non-fibrotic masks by intersecting each of those masks with the vessel and bronchial tubes one. The resulting masks are then also intersected with the polarization image to measure specifically the mean gray value and integrated intensity of fibrotic and non-fibrotic areas. The percentage of fibrotic area is obtained from the fibrotic and non-fibrotic masks without vessels and bronchial tubes as the ratio between the fibrotic area and the whole lung area without vessels and bronchial tubes (non-fibrotic + fibrotic).

Cells phenotyping by flow cytometry

Earlier version: Lungs or spleens were collected and digested with the lung or spleen dissociation kit (Miltenyi Biotech, ref 130-095-927 and 130-095-928) according to the supplier’s instructions. Red blood cells were lysed using ACK lysis buffer (Gibco, A1049201). Fc receptors were blocked using anti-CD 16/32 (2.4G2) antibodies followed by surface staining by incubating cells on ice, for 20 min, with saturating concentrations of labeled Abs (Table 1) in PBS, 5% FBS and 0.5% EDTA. Tregs were identified using the transcription factor staining Buffer Set (eBioscience) for FoxP3 staining. Intracellular staining was performed after stimulation of single-cell suspensions with Phorbol 12- myristate 13-acetate (PMA at 50 ng mL–1, Sigma), ionomycin (0.5 μg mL–1, Sigma) and 1 μL mL–1 Golgi Plug™ (BD Biosciences) for 4 h at 37°C 5% CO2. Cells were then incubated with Live/Dead stain (Invitrogen), according to the manufacturer protocol prior to surface staining. Fc receptors were blocked using anti-CD16/32 (2.4G2) antibodies followed by surface staining by incubating cells on ice, for 20 min, with saturating concentrations of labeled Abs (Table 1) in PBS, 5% FBS and 0.5% EDTA. Cytofix/Cytoperm™ kit (BD biosciences) was used for intracellular stainings by following the manufacturer’s instructions. Samples were acquired on CytoFLEX LX (Beckman Coulter) and analyzed using FlowJo (LLC).

Revised version: Lungs or spleens were collected and digested with the lung or spleen dissociation kit (Miltenyi Biotech, ref 130-095-927 and 130-095-928) according to the supplier’s instructions. Red blood cells were lysed using ACK lysis buffer (Gibco, A1049201). Fc receptors were blocked using anti-CD16/32 (2.4G2) antibodies followed by surface staining by incubating cells on ice, for 20 min, with saturating concentrations of labeled Abs (Table 1) in PBS, 5% FBS and 0.5% EDTA. Tregs were identified using the transcription factor staining Buffer Set (eBioscience) for FoxP3 staining. Intracellular staining was performed after stimulation of single-cell suspensions with Phorbol 12-myristate 13-acetate (PMA at 50 ng mL−1, Sigma), ionomycin (0.5 μg mL−1, Sigma) and 1 μL mL−1 Golgi Plug™ (BD Biosciences) for 4 h at 37°C 5% CO2. Cells were then incubated with Live/Dead stain (Invitrogen), according to the manufacturer protocol prior to surface staining. Fc receptors were blocked using anti-CD16/32 (2.4G2) antibodies followed by surface staining by incubating cells on ice, for 20 min, with saturating concentrations of labeled Abs (Table 1) in PBS, 5% FBS and 0.5% EDTA. Cytofix/Cytoperm™ kit (BD biosciences) was used for intracellular stainings by following the manufacturer’s instructions. Samples were acquired on CytoFLEX LX (Beckman Coulter) and analyzed using FlowJo (LLC).

ELISA

Sera of mice were collected at the end of the experiment and stored at −80 °C before cytokine detection by ELISA using mouse IL-1 beta/IL-1F2 (R&D, DY401) and IL-18 (MBL, ref 7625) according to the supplier’s instructions.

Western Blot

Earlier version: Single cell suspensions of whole lungs were lysed with Laemmli buffer (10% glycerol, 3% SDS, 10 mM Na2HPO4) with protease inhibitor cocktail (Roche). Proteins were separated on a 12% SDS-PAGE gel and electro transferred onto PVDF membranes, which were blocked for 30 min at RT with 3% bovine serum albumin or 5% milk. Membranes were incubated with primary antibodies (see Table 1) diluted at 4°C overnight. Secondary antibodies (Promega) were incubated for 1h at RT. Immunoblot detection was achieved by exposure with a chemiluminescence imaging system (PXI Syngene, Ozyme) after membrane incubation with ECL (Immobilon Western, Millipore). The bands intensity values were normalized to that of β-actin using ImageJ software.

Revised version: Single cell suspensions of whole lungs were lysed with Laemmli buffer (10% glycerol, 3% SDS, 10 mM Na2HPO4) with protease inhibitor cocktail (Roche). Proteins were separated on a 12% SDS-PAGE gel and electro transferred onto PVDF membranes, which were blocked for 30 min at RT with 3% bovine serum albumin or 5% milk. Membranes were incubated with primary antibodies (see Table 1) diluted at 4 °C overnight. Secondary antibodies (Promega) were incubated for 1 h at RT. Immunoblot detection was achieved by exposure with a chemiluminescence imaging system (PXI Syngene, Ozyme) after membrane incubation with ECL (Immobilon Western, Millipore). The bands intensity values were normalized to that of β-actin using ImageJ software.

Statistical analyses

Added in the revised version: Quantitative data were described and presented graphically as medians and interquartiles or means and standard deviations. The distribution normality was tested with the Shapiro’s test and homoscedasticity with a Bartlett’s test. For two categories, statistical comparisons were performed using the Student’s t-test or the Mann–Whitney’s test. For three and more categories, analysis of variance (ANOVA) or non-parametric data with Kruskal–Wallis was performed to test variables expressed as categories versus continuous variables. If this test was significant, we used the Tukey’s test to compare these categories and the Bonferroni’s test to adjust the significant threshold. For the Gene Set Enrichment Analyses (GSEA), bilateral Kolmogorov– Smirnov test, and false discovery rate (FDR) were used.

Added in the revised version: All statistical analyses were performed by biostatistician using Prism8 program from GraphPad software. Tests of significance was two-tailed and considered significant with an alpha level of P < 0.05. (graphically: * for P < 0.05, ** for P < 0.01, *** for P < 0.001).

Removed in the revised version: All analyses were carried out using Prism software (GraphPad). Data were obtained from at least n = 5 individuals. Mice were equally divided for treatments and controls. The presence of outlier was checked thanks to the ROULT method. Data is represented as mean values and error bars represent SEM. Two-tailed Mann–Whitney and unpaired t-test were used to evaluate the statistical significance between groups. For survival analyses, the log-rank Mantel-Cox test was used. For correlation analyses, Spearman test was used for the Gene Set Enrichment Analyses (GSEA).

Removed in the revised version: Data and materials availability

Removed in the revised version: All data are available in the main text or the supplementary materials. RNAseq data from IPF and control patients were retrieved from GEO database under the accession number (GSE47460, GPL14550).

Removed in the revised version: Author contributions

Removed in the revised version: Conceptualization and design: SJH, VV-C

Removed in the revised version: Methodology: SJH, FB

Removed in the revised version: Investigation: SJH, TJ, VV-C

Removed in the revised version: Analysis and Interpretation: SJH, SL, PH, VH, VV-C

Removed in the revised version: Reagents: AG

Removed in the revised version: Writing: SJH, VV-C